<section xml:id="exercises">
        <title>Exercises</title>
        <QuestionNode xml:id="functions-ex-1" line="14" names="functions_ex_1" runestone_options="{'number': 1, 'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': '', 'divid': 'functions_ex_1', 'qnumber': '', 'question_label': '6.13.1'}" source="/Users/bmiller/Runestone/books/thinkcspy/_sources/Functions/Exercises.rst">
            <TabbedStuffNode divid="q1" line="17" source="/Users/bmiller/Runestone/books/thinkcspy/_sources/Functions/Exercises.rst" tabbed_stuff_options="{'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': '', 'divid': 'q1', 'divclass': 'alert alert-warning'}">
                <TabNode tabname="Question" tabnode_options="{'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': '', 'tabname': 'Question'}">

<listing xml:id="ex_5_1">
    <caption/>
    <program xml:id="ex_5_1_editor" interactive="activecode" language="python">
        <input>

import turtle

def drawSquare(t, sz):
    """Get turtle t to draw a square of sz side"""

    for i in range(4):
        t.forward(sz)
        t.left(90)

wn = turtle.Screen()
wn.bgcolor("lightgreen")

alex = turtle.Turtle()
alex.color("pink")

drawSquare(alex,20)

wn.exitonclick()
        </input>
    </program>
</listing>
                    <p>Use the drawsquare function we wrote in this chapter in a program to draw
                        the image shown below.
                        Assume each side is 20 units.
                        (Hint: notice that the turtle has already moved away from the ending point of the last
                        square when the program ends.)</p>
                    <image source="images/Functions/Figures/five_squares.png"/>
                </TabNode>
                <TabNode tabname="Answer" tabnode_options="{'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': '', 'tabname': 'Answer'}">

<listing xml:id="func_q1_answer">
    <caption/>
    <program xml:id="func_q1_answer_editor" interactive="activecode" language="python">
        <input>
import turtle

def drawSquare(t, sz):
    """Make turtle t draw a square of with side sz."""
    for i in range(4):
        t.forward(sz)
        t.left(90)

wn = turtle.Screen()       # Set up the window and its attributes
wn.bgcolor("lightgreen")

alex = turtle.Turtle()     # create alex
alex.color('hotpink')
alex.pensize(3)

for i in range(5):
    drawSquare(alex, 20)   # Call the function to draw the square
    alex.penup()
    alex.forward(40)       # move alex to the starting position for the next square
    alex.pendown()

wn.exitonclick()
        </input>
    </program>
</listing>
                </TabNode>
                <TabNode tabname="Discussion" tabnode_options="{'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': '', 'tabname': 'Discussion'}">
                    <DisqusNode line="81" runestone_options="{'shortname': 'interactivepython', 'identifier': 'a2ac86a8d0524fc6830aefb785199048', 'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': ''}" source="/Users/bmiller/Runestone/books/thinkcspy/_sources/Functions/Exercises.rst">
                    </DisqusNode>
                </TabNode>
            </TabbedStuffNode>
        </QuestionNode>
        <QuestionNode xml:id="functions-ex-2" line="85" names="functions_ex_2" runestone_options="{'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': '', 'divid': 'functions_ex_2', 'qnumber': '', 'question_label': '6.13.4'}" source="/Users/bmiller/Runestone/books/thinkcspy/_sources/Functions/Exercises.rst">

<listing xml:id="ex_5_2">
    <caption/>
    <program xml:id="ex_5_2_editor" interactive="activecode" language="python">
        <input>

        </input>
    </program>
</listing>
            <p>Write a program to draw this. Assume the innermost square is 20 units per side,
                and each successive square is 20 units bigger, per side, than the one inside it.</p>
            <image source="images/Functions/Figures/nested_squares.png"/>
        </QuestionNode>
        <QuestionNode xml:id="functions-ex-3" line="96" names="functions_ex_3" runestone_options="{'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': '', 'divid': 'functions_ex_3', 'qnumber': '', 'question_label': '6.13.6'}" source="/Users/bmiller/Runestone/books/thinkcspy/_sources/Functions/Exercises.rst">
            <TabbedStuffNode divid="q3" line="98" source="/Users/bmiller/Runestone/books/thinkcspy/_sources/Functions/Exercises.rst" tabbed_stuff_options="{'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': '', 'divid': 'q3', 'divclass': 'alert alert-warning'}">
                <TabNode tabname="Question" tabnode_options="{'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': '', 'tabname': 'Question'}">

<listing xml:id="ex_5_3">
    <caption/>
    <program xml:id="ex_5_3_editor" interactive="activecode" language="python">
        <input>

        </input>
    </program>
</listing>
                    <p>Write a non-fruitful function <c>drawPoly(someturtle, somesides, somesize)</c> which makes a turtle
                        draw a regular polygon.
                        When called with <c>drawPoly(tess, 8, 50)</c>, it will draw a shape like this:</p>
                    <image source="images/Functions/Figures/regularpolygon.png"/>
                </TabNode>
                <TabNode tabname="Answer" tabnode_options="{'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': '', 'tabname': 'Answer'}">

<listing xml:id="func_q3_answer">
    <caption/>
    <program xml:id="func_q3_answer_editor" interactive="activecode" language="python">
        <input>
import turtle

def drawPoly(t, num_sides, side_length):
    for i in range(num_sides):
        t.forward(side_length)
        t.left(360/num_sides)

wn = turtle.Screen()       # Set up the window and its attributes
wn.bgcolor("lightgreen")

tess = turtle.Turtle()
tess.color('hotpink')
tess.pensize(3)

drawPoly(tess, 8, 50)
        </input>
    </program>
</listing>
                </TabNode>
                <TabNode tabname="Discussion" tabnode_options="{'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': '', 'tabname': 'Discussion'}">
                    <DisqusNode line="134" runestone_options="{'shortname': 'interactivepython', 'identifier': 'ba2f11265c524c7581bf7cf25d23bf3a', 'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': ''}" source="/Users/bmiller/Runestone/books/thinkcspy/_sources/Functions/Exercises.rst">
                    </DisqusNode>
                </TabNode>
            </TabbedStuffNode>
        </QuestionNode>
        <QuestionNode xml:id="functions-ex-4" line="138" names="functions_ex_4" runestone_options="{'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': '', 'divid': 'functions_ex_4', 'qnumber': '', 'question_label': '6.13.9'}" source="/Users/bmiller/Runestone/books/thinkcspy/_sources/Functions/Exercises.rst">

<listing xml:id="ex_5_4">
    <caption/>
    <program xml:id="ex_5_4_editor" interactive="activecode" language="python">
        <input>

        </input>
    </program>
</listing>
            <p>Draw this pretty pattern.</p>
            <image source="images/Functions/Figures/tess08.png"/>
        </QuestionNode>
        <QuestionNode xml:id="functions-ex-5" line="147" names="functions_ex_5" runestone_options="{'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': '', 'divid': 'functions_ex_5', 'qnumber': '', 'question_label': '6.13.11'}" source="/Users/bmiller/Runestone/books/thinkcspy/_sources/Functions/Exercises.rst">
            <TabbedStuffNode divid="q5" line="149" source="/Users/bmiller/Runestone/books/thinkcspy/_sources/Functions/Exercises.rst" tabbed_stuff_options="{'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': '', 'divid': 'q5', 'divclass': 'alert alert-warning'}">
                <TabNode tabname="Question" tabnode_options="{'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': '', 'tabname': 'Question'}">

<listing xml:id="ex_5_5">
    <caption/>
    <program xml:id="ex_5_5_editor" interactive="activecode" language="python">
        <input>

        </input>
    </program>
</listing>
                    <p>The two spirals in this picture differ only by the turn angle.  Draw both.</p>
                    <image source="images/Functions/Figures/tess_spirals.png" height="240"/>
                </TabNode>
                <TabNode tabname="Answer" tabnode_options="{'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': '', 'tabname': 'Answer'}">

<listing xml:id="func_q5_answer">
    <caption/>
    <program xml:id="func_q5_answer_editor" interactive="activecode" language="python">
        <input>
import turtle

def drawSpiral(t, angle):
    ''' takes a turtle, t, and an angle in degrees '''
    length = 1
    for i in range(84):
        t.forward(length)
        t.right(angle)
        length = length + 2


wn = turtle.Screen()       # Set up the window and its attributes
wn.bgcolor("lightgreen")

guido = turtle.Turtle()    # create guido
guido.color('blue')

## draw the first spiral ##
# position guido
guido.penup()
guido.backward(110)
guido.pendown()

# draw the spiral using a 90 degree turn angle
drawSpiral(guido, 90)


## draw the second spiral ##
# position guido
guido.home()
guido.penup()
guido.forward(90)
guido.pendown()

drawSpiral(guido, 89)
        </input>
    </program>
</listing>
                </TabNode>
                <TabNode tabname="Discussion" tabnode_options="{'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': '', 'tabname': 'Discussion'}">
                    <DisqusNode line="204" runestone_options="{'shortname': 'interactivepython', 'identifier': 'c587119991344db988f8fb37c8c9a31e', 'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': ''}" source="/Users/bmiller/Runestone/books/thinkcspy/_sources/Functions/Exercises.rst">
                    </DisqusNode>
                </TabNode>
            </TabbedStuffNode>
        </QuestionNode>
        <QuestionNode xml:id="functions-ex-6" line="208" names="functions_ex_6" runestone_options="{'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': '', 'divid': 'functions_ex_6', 'qnumber': '', 'question_label': '6.13.14'}" source="/Users/bmiller/Runestone/books/thinkcspy/_sources/Functions/Exercises.rst">

<listing xml:id="ex_5_6">
    <caption/>
    <program xml:id="ex_5_6_editor" interactive="activecode" language="python">
        <input>

        </input>
    </program>
</listing>
            <p>Write a non-fruitful function <c>drawEquitriangle(someturtle, somesize)</c> which calls <c>drawPoly</c> from the
                previous question to have its turtle draw a equilateral triangle.</p>
        </QuestionNode>
        <QuestionNode xml:id="functions-ex-7" line="217" names="functions_ex_7" runestone_options="{'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': '', 'divid': 'functions_ex_7', 'qnumber': '', 'question_label': '6.13.16'}" source="/Users/bmiller/Runestone/books/thinkcspy/_sources/Functions/Exercises.rst">
            <TabbedStuffNode divid="q7" line="219" source="/Users/bmiller/Runestone/books/thinkcspy/_sources/Functions/Exercises.rst" tabbed_stuff_options="{'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': '', 'divid': 'q7', 'divclass': 'alert alert-warning'}">
                <TabNode tabname="Question" tabnode_options="{'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': '', 'tabname': 'Question'}">

<listing xml:id="ex_5_7">
    <caption/>
    <program xml:id="ex_5_7_editor" interactive="activecode" language="python">
        <input>

def sumTo(n):
    # your code here

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertAlmostEqual(sumTo(15),120.0,0,"Tested sumTo on input 15")
        self.assertAlmostEqual(sumTo(0),0.0,0,"Tested sumTo on input 0")
        self.assertAlmostEqual(sumTo(25),325.0,0,"Tested sumTo on input 25")
        self.assertAlmostEqual(sumTo(7),28.0,0,"Tested sumTo on input 7")

myTests().main()
        </input>
    </program>
</listing>
                    <p>Write a fruitful function <c>sumTo(n)</c> that returns the sum of all integer numbers up to and
                        including <title_reference>n</title_reference>.   So <c>sumTo(10)</c> would be <c>1+2+3...+10</c> which would return the value 55.  Use the
                        equation  (n * (n + 1)) / 2.</p>
                </TabNode>
                <TabNode tabname="Answer" tabnode_options="{'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': '', 'tabname': 'Answer'}">

<listing xml:id="func_q7_answer">
    <caption/>
    <program xml:id="func_q7_answer_editor" interactive="activecode" language="python">
        <input>
from test import testEqual

def sumTo(n):
    result = (n * (n + 1)) / 2
    return result

# Now lets see how well this works
t = sumTo(0)
print("The sum from 1 to 0 is",t)
t = sumTo(10)
print("The sum from 1 to 10 is",t)
t = sumTo(5)
print("The sum from 1 to 5 is",t)
        </input>
    </program>
</listing>
                </TabNode>
                <TabNode tabname="Discussion" tabnode_options="{'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': '', 'tabname': 'Discussion'}">
                    <DisqusNode line="269" runestone_options="{'shortname': 'interactivepython', 'identifier': 'd6ba37a51d09845f39c96d4d4ef1d6f45', 'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': ''}" source="/Users/bmiller/Runestone/books/thinkcspy/_sources/Functions/Exercises.rst">
                    </DisqusNode>
                </TabNode>
            </TabbedStuffNode>
        </QuestionNode>
        <QuestionNode xml:id="functions-ex-8" line="273" names="functions_ex_8" runestone_options="{'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': '', 'divid': 'functions_ex_8', 'qnumber': '', 'question_label': '6.13.19'}" source="/Users/bmiller/Runestone/books/thinkcspy/_sources/Functions/Exercises.rst">

<listing xml:id="ex_5_8">
    <caption/>
    <program xml:id="ex_5_8_editor" interactive="activecode" language="python">
        <input>

def areaOfCircle(r):
    # your code here

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertAlmostEqual(areaOfCircle(5.0),78.53981633974483,5,"Tested input: areaOfCircle(5.0)")
        self.assertEqual(areaOfCircle(5.0),78.53981633974483,"Tested input: areaOfCirlce(5.0)")
        self.assertEqual(areaOfCircle(0),0.0,"Tested input: areaOfCirlce(0)")
        self.assertAlmostEqual(areaOfCircle(31415.926535897932),3100627668.0299816,5,"Tested input: areaOfCirlce(31415.926535897932)")


myTests().main()
        </input>
    </program>
</listing>
            <p>Write a function <title_reference>areaOfCircle(r)</title_reference> which returns the area of a circle of radius <title_reference>r</title_reference>.  Make sure you use the math module in your solution.</p>
        </QuestionNode>
        <QuestionNode xml:id="functions-ex-9" line="300" names="functions_ex_9" runestone_options="{'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': '', 'divid': 'functions_ex_9', 'qnumber': '', 'question_label': '6.13.21'}" source="/Users/bmiller/Runestone/books/thinkcspy/_sources/Functions/Exercises.rst">
            <TabbedStuffNode divid="q9" line="302" source="/Users/bmiller/Runestone/books/thinkcspy/_sources/Functions/Exercises.rst" tabbed_stuff_options="{'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': '', 'divid': 'q9', 'divclass': 'alert alert-warning'}">
                <TabNode tabname="Question" tabnode_options="{'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': '', 'tabname': 'Question'}">

<listing xml:id="ex_5_9">
    <caption/>
    <program xml:id="ex_5_9_editor" interactive="activecode" language="python">
        <input>

        </input>
    </program>
</listing>
                    <p>Write a non-fruitful function to draw a five pointed star, where the length of each side is 100 units.</p>
                    <image source="images/Functions/Figures/star.png"/>
                </TabNode>
                <TabNode tabname="Answer" tabnode_options="{'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': '', 'tabname': 'Answer'}">

<listing xml:id="func_q9_answer">
    <caption/>
    <program xml:id="func_q9_answer_editor" interactive="activecode" language="python">
        <input>
import turtle

def drawFivePointStar(t):
    for i in range(5):
        t.forward(100)
        t.left(216)

wolfram = turtle.Turtle()
drawFivePointStar(wolfram)
        </input>
    </program>
</listing>
                </TabNode>
                <TabNode tabname="Discussion" tabnode_options="{'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': '', 'tabname': 'Discussion'}">
                    <DisqusNode line="329" runestone_options="{'shortname': 'interactivepython', 'identifier': 'e757873187bb4581bffecdad449b5f61', 'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': ''}" source="/Users/bmiller/Runestone/books/thinkcspy/_sources/Functions/Exercises.rst">
                    </DisqusNode>
                </TabNode>
            </TabbedStuffNode>
        </QuestionNode>
        <QuestionNode xml:id="functions-ex-10" line="333" names="functions_ex_10" runestone_options="{'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': '', 'divid': 'functions_ex_10', 'qnumber': '', 'question_label': '6.13.24'}" source="/Users/bmiller/Runestone/books/thinkcspy/_sources/Functions/Exercises.rst">

<listing xml:id="ex_5_10">
    <caption/>
    <program xml:id="ex_5_10_editor" interactive="activecode" language="python">
        <input>

        </input>
    </program>
</listing>
            <p>Extend your program above.  Draw five stars, but between each, pick up the pen,
                move forward by 350 units, turn right by 144, put the pen down, and draw the next star.
                You&#8217;ll get something like this (note that you will need to move to the left before drawing your first star in order to fit everything in the window):</p>
            <image source="images/Functions/Figures/five_stars.png"/>
            <p>What would it look like if you didn&#8217;t pick up the pen?</p>
        </QuestionNode>
        <QuestionNode xml:id="functions-ex-11" line="347" names="functions_ex_11" runestone_options="{'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': '', 'divid': 'functions_ex_11', 'qnumber': '', 'question_label': '6.13.26'}" source="/Users/bmiller/Runestone/books/thinkcspy/_sources/Functions/Exercises.rst">
            <TabbedStuffNode divid="q11" line="349" source="/Users/bmiller/Runestone/books/thinkcspy/_sources/Functions/Exercises.rst" tabbed_stuff_options="{'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': '', 'divid': 'q11', 'divclass': 'alert alert-warning'}">
                <TabNode tabname="Question" tabnode_options="{'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': '', 'tabname': 'Question'}">

<listing xml:id="ex_5_11">
    <caption/>
    <program xml:id="ex_5_11_editor" interactive="activecode" language="python">
        <input>

        </input>
    </program>
</listing>
                    <p>Extend the star function to draw an n pointed star.  (Hint: n must be an odd number greater or
                        equal to 3).</p>
                </TabNode>
                <TabNode tabname="Answer" tabnode_options="{'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': '', 'tabname': 'Answer'}">

<listing xml:id="func_q11_answer">
    <caption/>
    <program xml:id="func_q11_answer_editor" interactive="activecode" language="python">
        <input>
import turtle

def drawStar(t, n):
    for i in range(n):
        t.forward(100)
        t.left(180 - 180/n)

stroustrup = turtle.Turtle()
drawStar(stroustrup, 7)
        </input>
    </program>
</listing>
                </TabNode>
                <TabNode tabname="Discussion" tabnode_options="{'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': '', 'tabname': 'Discussion'}">
                    <DisqusNode line="376" runestone_options="{'shortname': 'interactivepython', 'identifier': 'f2f8ff1b301e4d99bd4ac52e68c8c1ed', 'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': ''}" source="/Users/bmiller/Runestone/books/thinkcspy/_sources/Functions/Exercises.rst">
                    </DisqusNode>
                </TabNode>
            </TabbedStuffNode>
        </QuestionNode>
        <QuestionNode xml:id="functions-ex-12" line="380" names="functions_ex_12" runestone_options="{'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': '', 'divid': 'functions_ex_12', 'qnumber': '', 'question_label': '6.13.29'}" source="/Users/bmiller/Runestone/books/thinkcspy/_sources/Functions/Exercises.rst">

<listing xml:id="ex_5_12">
    <caption/>
    <program xml:id="ex_5_12_editor" interactive="activecode" language="python">
        <input>

        </input>
    </program>
</listing>
            <p>Write a function called drawSprite that will draw a sprite.  The function will need parameters for
                the turtle, the number of legs, and the length of the legs.  Invoke the function to create a sprite
                with 15 legs of length 120.</p>
        </QuestionNode>
        <QuestionNode xml:id="functions-ex-13" line="390" names="functions_ex_13" runestone_options="{'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': '', 'divid': 'functions_ex_13', 'qnumber': '', 'question_label': '6.13.31'}" source="/Users/bmiller/Runestone/books/thinkcspy/_sources/Functions/Exercises.rst">
            <TabbedStuffNode divid="q13" line="392" source="/Users/bmiller/Runestone/books/thinkcspy/_sources/Functions/Exercises.rst" tabbed_stuff_options="{'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': '', 'divid': 'q13', 'divclass': 'alert alert-warning'}">
                <TabNode tabname="Question" tabnode_options="{'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': '', 'tabname': 'Question'}">

<listing xml:id="ex_5_13">
    <caption/>
    <program xml:id="ex_5_13_editor" interactive="activecode" language="python">
        <input>

def sumTo(n):
    # your code here

====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(sumTo(15),120,"Tested sumTo on input 15")
        self.assertEqual(sumTo(0),0,"Tested sumTo on input 0")
        self.assertEqual(sumTo(25),325,"Tested sumTo on input 25")
        self.assertEqual(sumTo(7),28,"Tested sumTo on input 7")

myTests().main()
        </input>
    </program>
</listing>
                    <p>Rewrite the function <c>sumTo(n)</c> that returns the sum of all integer numbers up to and
                        including <title_reference>n</title_reference>.   This time use the accumulator pattern.</p>
                </TabNode>
                <TabNode tabname="Answer" tabnode_options="{'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': '', 'tabname': 'Answer'}">

<listing xml:id="func_q13_answer">
    <caption/>
    <program xml:id="func_q13_answer_editor" interactive="activecode" language="python">
        <input>
def sumTo(n):
    sum = 0
    for i in range(1,n+1):
        sum = sum + i
    return sum

# Now lets see how well this works
t = sumTo(0)
print("The sum from 1 to 0 is",t)
t = sumTo(10)
print("The sum from 1 to 10 is",t)
t = sumTo(5)
print("The sum from 1 to 5 is",t)
        </input>
    </program>
</listing>
                </TabNode>
                <TabNode tabname="Discussion" tabnode_options="{'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': '', 'tabname': 'Discussion'}">
                    <DisqusNode line="440" runestone_options="{'shortname': 'interactivepython', 'identifier': 'eda665389fda49a584b128cc30515595', 'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': ''}" source="/Users/bmiller/Runestone/books/thinkcspy/_sources/Functions/Exercises.rst">
                    </DisqusNode>
                </TabNode>
            </TabbedStuffNode>
        </QuestionNode>
        <QuestionNode xml:id="functions-ex-14" line="444" names="functions_ex_14" runestone_options="{'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': '', 'divid': 'functions_ex_14', 'qnumber': '', 'question_label': '6.13.34'}" source="/Users/bmiller/Runestone/books/thinkcspy/_sources/Functions/Exercises.rst">

<listing xml:id="ex_5_14">
    <caption/>
    <program xml:id="ex_5_14_editor" interactive="activecode" language="python">
        <input>

def mySqrt(n):
    # your code here

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
    def testOne(self):
        self.assertAlmostEqual(mySqrt(4.0),2.0,0,"Tested mySqrt on input 4.0")
        self.assertAlmostEqual(mySqrt(9.0),3.0,4,"Tested accuracy of mySqrt on input 3.0")
        self.assertAlmostEqual(mySqrt(36.0),6.0,5,"Tested accuracy of mySqrt on input 6.0")
        self.assertAlmostEqual(mySqrt(100.0),10.0,4,"Tested accuracy of mySqrt on input 10.0. Try iterating more times.")

myTests().main()
        </input>
    </program>
</listing>
            <p>Write a function called <c>mySqrt</c> that will approximate the square root of a number, call it n, by using
                Newton&#8217;s algorithm.
                Newton&#8217;s approach is an iterative guessing algorithm where the initial guess is n/2 and each subsequent guess
                is computed using   the formula:  newguess = (1/2) * (oldguess + (n/oldguess)).</p>
        </QuestionNode>
        <QuestionNode xml:id="functions-ex-15" line="472" names="functions_ex_15" runestone_options="{'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': '', 'divid': 'functions_ex_15', 'qnumber': '', 'question_label': '6.13.36'}" source="/Users/bmiller/Runestone/books/thinkcspy/_sources/Functions/Exercises.rst">
            <TabbedStuffNode divid="q15" line="474" source="/Users/bmiller/Runestone/books/thinkcspy/_sources/Functions/Exercises.rst" tabbed_stuff_options="{'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': '', 'divid': 'q15', 'divclass': 'alert alert-warning'}">
                <TabNode tabname="Question" tabnode_options="{'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': '', 'tabname': 'Question'}">

<listing xml:id="ex_5_15">
    <caption/>
    <program xml:id="ex_5_15_editor" interactive="activecode" language="python">
        <input>

 def myPi(iters):
     # Calculate an approximation of PI using the Leibniz
     # approximation with iters number of iterations

     # your code here
        </input>
    </program>
</listing>
                    <p>Write a function called <c>myPi</c> that will return an approximation of PI (3.14159&#8230;).  Use the <url href="http://en.wikipedia.org/wiki/Leibniz_formula_for_%CF%80" visual="http://en.wikipedia.org/wiki/Leibniz_formula_for_%CF%80">Leibniz</url> approximation.</p>
                </TabNode>
                <TabNode tabname="Answer" tabnode_options="{'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': '', 'tabname': 'Answer'}">

<listing xml:id="func_q15_answer">
    <caption/>
    <program xml:id="func_q15_answer_editor" interactive="activecode" language="python">
        <input>
def myPi(iters):
    ''' Calculate an approximation of PI using the Leibniz
    approximation with iters number of iterations '''
    pi = 0
    sign = 1
    denominator = 1
    for i in range(iters):
        pi = pi + (sign/denominator)
        sign = sign * -1  # alternate positive and negative
        denominator = denominator + 2

    pi = pi * 4.0
    return pi

pi_approx = myPi(10000)
print(pi_approx)
        </input>
    </program>
</listing>
                </TabNode>
                <TabNode tabname="Discussion" tabnode_options="{'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': '', 'tabname': 'Discussion'}">
                    <DisqusNode line="513" runestone_options="{'shortname': 'interactivepython', 'identifier': 'b699e4b7bad44db6bd788c795c124b23', 'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': ''}" source="/Users/bmiller/Runestone/books/thinkcspy/_sources/Functions/Exercises.rst">
                    </DisqusNode>
                </TabNode>
            </TabbedStuffNode>
        </QuestionNode>
        <QuestionNode xml:id="functions-ex-16" line="517" names="functions_ex_16" runestone_options="{'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': '', 'divid': 'functions_ex_16', 'qnumber': '', 'question_label': '6.13.39'}" source="/Users/bmiller/Runestone/books/thinkcspy/_sources/Functions/Exercises.rst">

<listing xml:id="ex_5_16">
    <caption/>
    <program xml:id="ex_5_16_editor" interactive="activecode" language="python">
        <input>

def myPi(iters):
    # Calculate an approximation of PI using the Madhava
    # approximation with iters number of iterations

    #your code here
        </input>
    </program>
</listing>
            <p>Write a function called <title_reference>myPi</title_reference> that will return an approximation of PI (3.14159&#8230;).  Use the <url href="http://en.wikipedia.org/wiki/Madhava_of_Sangamagrama" visual="http://en.wikipedia.org/wiki/Madhava_of_Sangamagrama">Madhava</url> approximation.</p>
        </QuestionNode>
        <QuestionNode xml:id="functions-ex-17" line="530" names="functions_ex_17" runestone_options="{'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': '', 'divid': 'functions_ex_17', 'qnumber': '', 'question_label': '6.13.41'}" source="/Users/bmiller/Runestone/books/thinkcspy/_sources/Functions/Exercises.rst">
            <TabbedStuffNode divid="q17" line="532" source="/Users/bmiller/Runestone/books/thinkcspy/_sources/Functions/Exercises.rst" tabbed_stuff_options="{'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': '', 'divid': 'q17', 'divclass': 'alert alert-warning'}">
                <TabNode tabname="Question" tabnode_options="{'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': '', 'tabname': 'Question'}">

<listing xml:id="ex_5_17">
    <caption/>
    <program xml:id="ex_5_17_editor" interactive="activecode" language="python">
        <input>

        </input>
    </program>
</listing>
                    <p>Write a function called <title_reference>fancySquare</title_reference> that will draw a square with fancy corners (sprites on the corners).  You should
                        implement and use the <title_reference>drawSprite</title_reference> function from above.  For an even more interesting look, how about adding small
                        triangles to the ends of the sprite legs.</p>
                </TabNode>
                <TabNode tabname="Answer" tabnode_options="{'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': '', 'tabname': 'Answer'}">

<listing xml:id="func_q17_answer">
    <caption/>
    <program xml:id="func_q17_answer_editor" interactive="activecode" language="python">
        <input>
import turtle

def drawSprite(t, numlegs, leglength):
   angle = 360/numlegs
   for i in range(numlegs):
      t.forward(leglength)
      t.backward(leglength)
      t.left(angle)

def drawFancySquare(t, sz, lgs, lgl):
   for i in range(4):
       t.forward(sz)
       drawSprite(t, lgs, lgl)
       t.left(90)

wn = turtle.Screen()
wn.bgcolor("lightgreen")

alex = turtle.Turtle()
drawFancySquare(alex, 100, 10, 15)

wn.exitonclick()
        </input>
    </program>
</listing>
                </TabNode>
                <TabNode tabname="Discussion" tabnode_options="{'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': '', 'tabname': 'Discussion'}">
                    <DisqusNode line="572" runestone_options="{'shortname': 'interactivepython', 'identifier': 'db5d8808bf5749579718bdd2088b539f', 'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': ''}" source="/Users/bmiller/Runestone/books/thinkcspy/_sources/Functions/Exercises.rst">
                    </DisqusNode>
                </TabNode>
            </TabbedStuffNode>
        </QuestionNode>
        <QuestionNode xml:id="functions-ex-18" line="576" names="functions_ex_18" runestone_options="{'subchapter': 'Exercises', 'chapter': 'Functions', 'basecourse': 'thinkcspy', 'optional': '', 'optclass': '', 'divid': 'functions_ex_18', 'qnumber': '', 'question_label': '6.13.44'}" source="/Users/bmiller/Runestone/books/thinkcspy/_sources/Functions/Exercises.rst">

<listing xml:id="ex_5_18">
    <caption/>
    <program xml:id="ex_5_18_editor" interactive="activecode" language="python">
        <input>

import turtle

def drawBar(t, height):
    """ Get turtle t to draw one bar, of height. """
    t.begin_fill()               # start filling this shape
    t.left(90)
    t.forward(height)
    t.write(str(height))
    t.right(90)
    t.forward(40)
    t.right(90)
    t.forward(height)
    t.left(90)
    t.end_fill()                 # stop filling this shape



xs = [48, 117, 200, 240, 160, 260, 220]  # here is the data
maxheight = max(xs)
numbars = len(xs)
border = 10

wn = turtle.Screen()             # Set up the window and its attributes
wn.setworldcoordinates(0-border, 0-border, 40*numbars+border, maxheight+border)
wn.bgcolor("lightgreen")

tess = turtle.Turtle()           # create tess and set some attributes
tess.color("blue")
tess.fillcolor("red")
tess.pensize(3)



for a in xs:
    drawBar(tess, a)

wn.exitonclick()
        </input>
    </program>
</listing>
            <p>There was a whole program in <xref ref="ATurtleBarChart#bar-chart"/> to create a bar chart with specific data.  Creating a bar chart is a useful idea in general.  Write a non-fruitful function called barChart, that takes the numeric list of data as a parameter, and draws the bar chart.  Write a full program calling this function.
                The current version of the <c>drawBar</c> function unfortuately draws the top of the bar through the bottom of the label.  A nice elaboration is to make the label appear completely above the top line.  To keep the spacing consistent you might pass an extra parameter to <c>drawBar</c> for the distance to move up.  For the <c>barChart</c> function make that parameter be some small fraction of <c>maxheight+border</c>.  The fill action makes this modification particularly tricky:  You will want to move past the top of the bar and write before or after drawing and filling the bar.</p>
        </QuestionNode>
    </section>

