<!DOCTYPE html>


<html xmlns="http://www.w3.org/1999/xhtml">
  <head>
    <meta http-equiv="Content-Type" content="text/html; charset=utf-8" />
    
    <title>Exercises &mdash; How to Think like a Computer Scientist: Interactive Edition</title>
    
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../_static/bootstrap-3.0.0/css/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/tabbedstuff.css" type="text/css" />
    <link rel="stylesheet" href="../_static/pytutor.css" type="text/css" />
    <link rel="stylesheet" href="../_static/modal-basic.css" type="text/css" />
    <link rel="stylesheet" href="../_static/clickable.css" type="text/css" />
    <link rel="stylesheet" href="../_static/codemirror.css" type="text/css" />
    <link rel="stylesheet" href="../_static/livecode.css" type="text/css" />
    <link rel="stylesheet" href="../_static/video.css" type="text/css" />
    <link rel="stylesheet" href="../_static/fitb.css" type="text/css" />
    <link rel="stylesheet" href="../_static/datafile.css" type="text/css" />
    <link rel="stylesheet" href="../_static/codemirror.css" type="text/css" />
    <link rel="stylesheet" href="../_static/activecode.css" type="text/css" />
    <link rel="stylesheet" href="../_static/parsons.css" type="text/css" />
    <link rel="stylesheet" href="../_static/lib/prettify.css" type="text/css" />
    <link rel="stylesheet" href="../_static/poll.css" type="text/css" />
    <link rel="stylesheet" href="../_static/dragndrop.css" type="text/css" />
    <link rel="stylesheet" href="http://netdna.bootstrapcdn.com/bootswatch/2.3.1/""/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/jquery-ui-1.10.3.custom.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/bootstrap-sphinx.css" type="text/css" />
    <link rel="stylesheet" href="../_static/user-highlights.css" type="text/css" />
    <link rel="stylesheet" href="../_static/runestone-custom-sphinx-bootstrap.css" type="text/css" />
    
    <script type="text/javascript">
      var DOCUMENTATION_OPTIONS = {
        URL_ROOT:    '../',
        VERSION:     '3.0',
        COLLAPSE_INDEX: false,
        FILE_SUFFIX: '.html',
        HAS_SOURCE:  true
      };
    </script>
    <script type="text/javascript" src="../_static/runestonebase.js"></script>
    <script type="text/javascript" src="../_static/jquery.js"></script>
    <script type="text/javascript" src="../_static/underscore.js"></script>
    <script type="text/javascript" src="../_static/doctools.js"></script>
    <script type="text/javascript" src="../_static/tabbedstuff.js"></script>
    <script type="text/javascript" src="../_static/animationbase.js"></script>
    <script type="text/javascript" src="../_static/d3.v2.min.js"></script>
    <script type="text/javascript" src="../_static/jquery.ba-bbq.min.js"></script>
    <script type="text/javascript" src="../_static/jquery.jsPlumb-1.3.10-all-min.js"></script>
    <script type="text/javascript" src="../_static/pytutor.js"></script>
    <script type="text/javascript" src="../_static/codelens.js"></script>
    <script type="text/javascript" src="../_static/clickable.js"></script>
    <script type="text/javascript" src="../_static/timedclickable.js"></script>
    <script type="text/javascript" src="../_static/livecode.js"></script>
    <script type="text/javascript" src="../_static/clike.js"></script>
    <script type="text/javascript" src="../_static/mchoice.js"></script>
    <script type="text/javascript" src="../_static/timedmc.js"></script>
    <script type="text/javascript" src="../_static/fitb.js"></script>
    <script type="text/javascript" src="../_static/timedfitb.js"></script>
    <script type="text/javascript" src="../_static/timed.js"></script>
    <script type="text/javascript" src="../_static/skulpt.min.js"></script>
    <script type="text/javascript" src="../_static/skulpt-stdlib.js"></script>
    <script type="text/javascript" src="../_static/datafile.js"></script>
    <script type="text/javascript" src="../_static/jquery.highlight.js"></script>
    <script type="text/javascript" src="../_static/bookfuncs.js"></script>
    <script type="text/javascript" src="../_static/codemirror.js"></script>
    <script type="text/javascript" src="../_static/xml.js"></script>
    <script type="text/javascript" src="../_static/css.js"></script>
    <script type="text/javascript" src="../_static/htmlmixed.js"></script>
    <script type="text/javascript" src="../_static/python.js"></script>
    <script type="text/javascript" src="../_static/javascript.js"></script>
    <script type="text/javascript" src="../_static/activecode.js"></script>
    <script type="text/javascript" src="../_static/skulpt.min.js"></script>
    <script type="text/javascript" src="../_static/skulpt-stdlib.js"></script>
    <script type="text/javascript" src="../_static/clike.js"></script>
    <script type="text/javascript" src="../_static/timed_activecode.js"></script>
    <script type="text/javascript" src="../_static/lib/jquery.min.js"></script>
    <script type="text/javascript" src="../_static/lib/jquery-ui.min.js"></script>
    <script type="text/javascript" src="../_static/lib/jquery.ui.touch-punch.min.js"></script>
    <script type="text/javascript" src="../_static/lib/prettify.js"></script>
    <script type="text/javascript" src="../_static/lib/underscore-min.js"></script>
    <script type="text/javascript" src="../_static/lib/lis.js"></script>
    <script type="text/javascript" src="../_static/parsons_setup.js"></script>
    <script type="text/javascript" src="../_static/parsons.js"></script>
    <script type="text/javascript" src="../_static/parsons-noconflict.js"></script>
    <script type="text/javascript" src="../_static/timedparsons.js"></script>
    <script type="text/javascript" src="../_static/shortanswer.js"></script>
    <script type="text/javascript" src="../_static/timed_shortanswer.js"></script>
    <script type="text/javascript" src="../_static/reveal.js"></script>
    <script type="text/javascript" src="../_static/poll.js"></script>
    <script type="text/javascript" src="../_static/dragndrop.js"></script>
    <script type="text/javascript" src="../_static/timeddnd.js"></script>
    <script type="text/javascript" src="https://cdn.mathjax.org/mathjax/latest/MathJax.js?config=TeX-AMS-MML_HTMLorMML"></script>
    <script type="text/javascript" src="../_static/jquery-ui-1.10.3.custom.min.js"></script>
    <script type="text/javascript" src="../_static/jquery-fix.js"></script>
    <script type="text/javascript" src="../_static/bootstrap-3.0.0/js/bootstrap.min.js"></script>
    <script type="text/javascript" src="../_static/bootstrap-sphinx.js"></script>
    <script type="text/javascript" src="../_static/waypoints.min.js"></script>
    <script type="text/javascript" src="../_static/rangy-core.js"></script>
    <script type="text/javascript" src="../_static/rangy-textrange.js"></script>
    <script type="text/javascript" src="../_static/rangy-cssclassapplier.js"></script>
    <script type="text/javascript" src="../_static/user-highlights.js"></script>
    <script type="text/javascript" src="../_static/jquery.idle-timer.js"></script>
    <script type="text/javascript" src="../_static/processing-1.4.1.min.js"></script>
    <script type="text/javascript" src="../_static/jquery.hotkey.js"></script>
    <script type="text/javascript" src="../_static/jquery-migrate-1.2.1.min.js"></script>
    <link rel="top" title="How to Think like a Computer Scientist: Interactive Edition" href="../toc.html" />
    <link rel="next" title="Strings Revisited" href="../Strings/StringsRevisited.html" />
    <link rel="prev" title="Glossary" href="Glossary.html" />
<meta charset='utf-8'>
<meta http-equiv='X-UA-Compatible' content='IE=edge,chrome=1'>
<meta content='width=device-width, initial-scale=1.0, maximum-scale=1.0, user-scalable=0' name='viewport' />
<script type="text/javascript">
  eBookConfig = {};
  eBookConfig.host = 'http://127.0.0.1:8000' ? 'http://127.0.0.1:8000' : 'http://127.0.0.1:8000';
  eBookConfig.app = eBookConfig.host+'/runestone';
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.course = 'thinkcspy';
  eBookConfig.logLevel = 10;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "3.5.1";
  eBookConfig.isLoggedIn = false;
  eBookConfig.useRunestoneServices = true;
  eBookConfig.python3 = true;
  eBookConfig.basecourse = 'thinkcspy';
</script>

<div id="fb-root"></div>
<script>
  (function(d, s, id) {
    var js, fjs = d.getElementsByTagName(s)[0];
    if (d.getElementById(id)) return;
    js = d.createElement(s); js.id = id;
    js.src = "//connect.facebook.net/en_US/all.js#xfbml=1&status=0";
    fjs.parentNode.insertBefore(js, fjs);
  }(document, 'script', 'facebook-jssdk'));
</script>



  </head>
  <body role="document">


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse">
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        <a class="brand-logo" style='background: transparent url("../_static/img/logo_small.png") no-repeat 0px 0px;' href="/runestone/default/user/login">&nbsp; </a>
        <a class="navbar-brand" href="../toc.html">How To Think Like a Computer Scientist</a>
      </div>
    </div>

    <div class="navbar-collapse collapse navbar-ex1-collapse">

      <ul class="nav navbar-nav navbar-right">

        <li class="divider-vertical"></li>

        <!-- social media dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-share-alt" style="opacity: 0.9"></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone" class="twitter-follow-button" data-show-count="true">Follow @iRunestone</a><script>!function(d,s,id){var js,fjs=d.getElementsByTagName(s)[0],p=/^http:/.test(d.location)?'http':'https';if(!d.getElementById(id)){js=d.createElement(s);js.id=id;js.src=p+'://platform.twitter.com/widgets.js';fjs.parentNode.insertBefore(js,fjs);}}(document, 'script', 'twitter-wjs');</script>
                <div class="fb-like" data-href="https://www.facebook.com/RunestoneInteractive" data-send="false" data-layout="button_count" data-width="300" data-show-faces="false" data-font="arial"></div>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div class="gittip">
                  <iframe style="border:0;"
                    src="https://www.gittip.com/bnmnetp/widget.html"
                    width="48pt" height="22pt" class="gittip-button">
                  </iframe>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li>
        <!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../toc.html'>Table of Contents</a></li>
            
            <li><a href='/runestone/static/thinkcspy/genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        <!-- user account dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-user" style="opacity:0.9;"></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><span class='loggedinuser'></span></li>
            <li class="divider"></li>
            <li><a href='http://127.0.0.1:8000/runestone/mygroup/initiateGroup' id="joinGroupLink">Join a Study Group</a></li>
            <li><a href='http://127.0.0.1:8000/runestone/mygroup/schedule' id="groupScheduleLink">Group Schedule</a></li>
            <li><a href='http://127.0.0.1:8000/runestone/mygroup/newschedule' id="newChapterLink">Schedule New Chapter</a></li>
            <li><a href='http://127.0.0.1:8000/runestone/mygroup/manageGroup' id="manageGroupLink">Manage Group</a></li>
			<li class="divider"></li>
            <li><a href='http://127.0.0.1:8000/runestone/admin/index'>Instructor's Page</a></li>
            <li><a href='http://127.0.0.1:8000/runestone/assignments/index'>Progress Page</a></li>
            <li class="divider"></li>
            <li><a href="http://127.0.0.1:8000/runestone/default/user/profile" id="profilelink">Edit Profile</a></li>
            <li><a href="http://127.0.0.1:8000/runestone/default/user/change_password" id="passwordlink">Change Password</a></li>
            <li><a href="http://127.0.0.1:8000/runestone/default/user/register" id="registerlink">Register</a></li>
            <li class='loginout'><a href='#'>Login</a></li> <!-- correct link populated by addNavbarLoginLink() -->
          </ul>
        </li>
        <!-- end user account dropdown -->

        <li class="divider-vertical"></li>

        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='/runestone/static/thinkcspy/navhelp.html'>Navigation Help</a></li>
            <li><a href='/runestone/static/overview/instructor.html'>Help for Instructors</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/static/help/bugreports.html'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>

      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        <!--
          <li class="dropdown">
  <a href="#" class="dropdown-toggle" data-toggle="dropdown">Page <b class="caret"></b></a>
  <ul class="dropdown-menu localtoc"><ul>
<li><a class="reference internal" href="#">Exercises</a></li>
</ul>
</ul>
</li>
          <li class="divider-vertical"></li>
        
        
          
  <li id="relations-prev" title="Previous Chapter - Glossary" data-toggle="tooltip">
    <a href="Glossary.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - Strings Revisited' data-toggle="tooltip" >
    <a href="../Strings/StringsRevisited.html" >
      <i class='glyphicon glyphicon-forward' style='opacity:0.9;'></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script type="text/javascript">
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  $('#relations-prev').tooltip(opts);
  $('#relations-next').tooltip(opts);
</script>
        -->
        
          <li></li>
        
      </ul>

    </div>
  </div>
</div>


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content">
  
  <div class="section" id="exercises">
<h1>Exercises<a class="headerlink" href="#exercises" title="Permalink to this headline">¶</a></h1>
<div class="full-width container">
<ol class="arabic">
<li><div id='q1' data-component="tabbedStuff" ><div data-component="tab" data-tabname="Question" >
<p>Add a print statement to Newton&#8217;s <code class="docutils literal"><span class="pre">sqrt</span></code> function that
prints out <code class="docutils literal"><span class="pre">better</span></code> each time it is calculated. Call your modified
function with 25 as an argument and record the results.</p>

<textarea data-component="activecode" id=ex_7_7 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true" data-audio=''    data-gradebutton=true >
def newtonSqrt(n):
    # your code here
</textarea>
</div><div data-component="tab" data-tabname="Answer" >

<textarea data-component="activecode" id=q1_answer data-lang="python"    data-timelimit=25000  data-codelens="true" data-audio=''     >
def newtonSqrt(n):
    approx = 0.5 * n
    better = 0.5 * (approx + n/approx)
    while better != approx:
        approx = better
        better = 0.5 * (approx + n/approx)
        print("Approx:", better)
    return approx


print("Final approx:", newtonSqrt(25))
</textarea>
</div><div data-component="tab" data-tabname="Discussion" >
<script type="text/javascript">
    function disqus_5784e08291ba43199d43fdab277849f5_disqus(source) { 
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'disqus_5784e08291ba43199d43fdab277849f5_' + eBookConfig.course;
                    this.page.url = 'http://disqus_5784e08291ba43199d43fdab277849f5_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';    
            disqus_identifier = 'disqus_5784e08291ba43199d43fdab277849f5_' + eBookConfig.course;
            disqus_url = 'http://disqus_5784e08291ba43199d43fdab277849f5_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'http://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class='disqus_thread_link' data-disqus-identifier="disqus_5784e08291ba43199d43fdab277849f5" onclick="disqus_5784e08291ba43199d43fdab277849f5_disqus(this);">Show Comments</a>
<script type='text/javascript'>
    $("a[data-disqus-identifier='disqus_5784e08291ba43199d43fdab277849f5']").attr('data-disqus-identifier', 'disqus_5784e08291ba43199d43fdab277849f5_' + eBookConfig.course);
</script>
</div>
    </div>

</li>
<li><p class="first">Write a function <code class="docutils literal"><span class="pre">print_triangular_numbers(n)</span></code> that prints out the first
n triangular numbers. A call to <code class="docutils literal"><span class="pre">print_triangular_numbers(5)</span></code> would
produce the following output:</p>
<div class="highlight-default"><div class="highlight"><pre><span class="mi">1</span>       <span class="mi">1</span>
<span class="mi">2</span>       <span class="mi">3</span>
<span class="mi">3</span>       <span class="mi">6</span>
<span class="mi">4</span>       <span class="mi">10</span>
<span class="mi">5</span>       <span class="mi">15</span>
</pre></div>
</div>
<p>(<em>hint: use a web search to find out what a triangular number is.</em>)</p>

<textarea data-component="activecode" id=ex_7_8 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true" data-audio=''    data-gradebutton=true >
def print_triangular_numbers(n):
    # your code here
</textarea>
</li>
<li><div id='q3' data-component="tabbedStuff" ><div data-component="tab" data-tabname="Question" >
<p>Write a function, <code class="docutils literal"><span class="pre">is_prime</span></code>, that takes a single integer argument
and returns <code class="docutils literal"><span class="pre">True</span></code> when the argument is a <em>prime number</em> and <code class="docutils literal"><span class="pre">False</span></code>
otherwise.</p>

<textarea data-component="activecode" id=ex_7_9 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true" data-audio=''    data-gradebutton=true >
def is_prime(n):
    # your code here

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(is_prime(2),True,"2 is a prime number.")
        self.assertEqual(is_prime(4187),False,"4187 is not a prime number, it is divisible by 53 and 79.")
        self.assertEqual(is_prime(22),False,"22 is not a prime number, it is divisible by 2 and 11.")
        self.assertEqual(is_prime(4813),True,"4813 is a prime number.")

myTests().main()
</textarea>
</div><div data-component="tab" data-tabname="Answer" >

<textarea data-component="activecode" id=q3_answer data-lang="python"    data-timelimit=25000  data-codelens="true" data-audio=''     >
def is_prime(n):
    for i in range(2, n):
        if n % i == 0:
            return False
    return True

print(is_prime(25))
print(is_prime(7))
print(is_prime(251))
print(is_prime(20))
</textarea>
</div><div data-component="tab" data-tabname="Discussion" >
<script type="text/javascript">
    function disqus_418de05233374e76b3b66aeb96b55656_disqus(source) { 
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'disqus_418de05233374e76b3b66aeb96b55656_' + eBookConfig.course;
                    this.page.url = 'http://disqus_418de05233374e76b3b66aeb96b55656_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';    
            disqus_identifier = 'disqus_418de05233374e76b3b66aeb96b55656_' + eBookConfig.course;
            disqus_url = 'http://disqus_418de05233374e76b3b66aeb96b55656_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'http://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class='disqus_thread_link' data-disqus-identifier="disqus_418de05233374e76b3b66aeb96b55656" onclick="disqus_418de05233374e76b3b66aeb96b55656_disqus(this);">Show Comments</a>
<script type='text/javascript'>
    $("a[data-disqus-identifier='disqus_418de05233374e76b3b66aeb96b55656']").attr('data-disqus-identifier', 'disqus_418de05233374e76b3b66aeb96b55656_' + eBookConfig.course);
</script>
</div>
    </div>

</li>
<li><p class="first">Modify the walking turtle program so that rather than a 90 degree left or right turn the
angle of the turn is determined randomly at each step.</p>
<blockquote>
<div>
<textarea data-component="activecode" id=ex_7_14 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true"  data-audio=''    data-gradebutton=true >


</textarea>
</div></blockquote>
</li>
<li><div id='q5' data-component="tabbedStuff" ><div data-component="tab" data-tabname="Question" >
<p>Modify the turtle walk program so that you have two turtles each with a
random starting location.  Keep the turtles moving until one of them leaves the screen.</p>

<textarea data-component="activecode" id=ex_7_13 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true"  data-audio=''    data-gradebutton=true >


</textarea>
</div><div data-component="tab" data-tabname="Answer" >

<textarea data-component="activecode" id=q5_answer data-lang="python"    data-timelimit=25000   data-audio=''     >
import random
import turtle

def moveRandom(wn, t):
    coin = random.randrange(0,2)
    if coin == 0:
        t.left(90)
    else:
        t.right(90)

    t.forward(50)

def areColliding(t1, t2):
    if t1.distance(t2) < 2:
        return True
    else:
        return False

def isInScreen(w, t):
    leftBound = - w.window_width() / 2
    rightBound = w.window_width() / 2
    topBound = w.window_height() / 2
    bottomBound = -w.window_height() / 2

    turtleX = t.xcor()
    turtleY = t.ycor()

    stillIn = True
    if turtleX > rightBound or turtleX < leftBound:
        stillIn = False
    if turtleY > topBound or turtleY < bottomBound:
        stillIn = False
    return stillIn

t1 = turtle.Turtle()
t2 = turtle.Turtle()
wn = turtle.Screen()

t1.shape('turtle')
t2.shape('circle')

leftBound = -wn.window_width() / 2
rightBound = wn.window_width() / 2
topBound = wn.window_height() / 2
bottomBound = -wn.window_height() / 2

t1.up()
t1.goto(random.randrange(leftBound, rightBound),
        random.randrange(bottomBound, topBound))
t1.setheading(random.randrange(0, 360))
t1.down()

t2.up()
t2.goto(random.randrange(leftBound, rightBound),
        random.randrange(bottomBound, topBound))
t2.setheading(random.randrange(0, 360))
t2.down()


while isInScreen(wn, t1) and isInScreen(wn, t2):
    moveRandom(wn, t1)
    moveRandom(wn, t2)

wn.exitonclick()
</textarea>
</div><div data-component="tab" data-tabname="Discussion" >
<script type="text/javascript">
    function disqus_0cd01637a1814f86b11f576c37a46437_disqus(source) { 
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'disqus_0cd01637a1814f86b11f576c37a46437_' + eBookConfig.course;
                    this.page.url = 'http://disqus_0cd01637a1814f86b11f576c37a46437_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';    
            disqus_identifier = 'disqus_0cd01637a1814f86b11f576c37a46437_' + eBookConfig.course;
            disqus_url = 'http://disqus_0cd01637a1814f86b11f576c37a46437_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'http://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class='disqus_thread_link' data-disqus-identifier="disqus_0cd01637a1814f86b11f576c37a46437" onclick="disqus_0cd01637a1814f86b11f576c37a46437_disqus(this);">Show Comments</a>
<script type='text/javascript'>
    $("a[data-disqus-identifier='disqus_0cd01637a1814f86b11f576c37a46437']").attr('data-disqus-identifier', 'disqus_0cd01637a1814f86b11f576c37a46437_' + eBookConfig.course);
</script>
</div>
    </div>

</li>
<li><p class="first">Modify the previous turtle walk program so that the turtle turns around
when it hits the wall or when one turtle collides with another turtle.</p>

<textarea data-component="activecode" id=ex_7_12 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true"  data-audio=''    data-gradebutton=true >


</textarea>
</li>
<li><div id='q7' data-component="tabbedStuff" ><div data-component="tab" data-tabname="Question" >
<p>Write a function to remove all the red from an image.</p>
<img src="../_static/LutherBellPic.jpg" id="luther.jpg">
<h4 style="text-align: left;">For this and the following exercises, use the
luther.jpg photo.</h4>
<textarea data-component="activecode" id=ex_7_15 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true"  data-audio=''    data-gradebutton=true >


</textarea>
</div><div data-component="tab" data-tabname="Answer" >

<textarea data-component="activecode" id=q7_answer data-lang="python"    data-timelimit=25000   data-audio=''     >
import image

img = image.Image("luther.jpg")
newimg = image.EmptyImage(img.getWidth(), img.getHeight())
win = image.ImageWin()

for col in range(img.getWidth()):
    for row in range(img.getHeight()):
        p = img.getPixel(col, row)

        newred = 0
        green = p.getGreen()
        blue = p.getBlue()

        newpixel = image.Pixel(newred, green, blue)

        newimg.setPixel(col, row, newpixel)

newimg.draw(win)
win.exitonclick()
</textarea>
</div><div data-component="tab" data-tabname="Discussion" >
<script type="text/javascript">
    function disqus_777006b154ca4af7ab8bd11cc25c208a_disqus(source) { 
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'disqus_777006b154ca4af7ab8bd11cc25c208a_' + eBookConfig.course;
                    this.page.url = 'http://disqus_777006b154ca4af7ab8bd11cc25c208a_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';    
            disqus_identifier = 'disqus_777006b154ca4af7ab8bd11cc25c208a_' + eBookConfig.course;
            disqus_url = 'http://disqus_777006b154ca4af7ab8bd11cc25c208a_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'http://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class='disqus_thread_link' data-disqus-identifier="disqus_777006b154ca4af7ab8bd11cc25c208a" onclick="disqus_777006b154ca4af7ab8bd11cc25c208a_disqus(this);">Show Comments</a>
<script type='text/javascript'>
    $("a[data-disqus-identifier='disqus_777006b154ca4af7ab8bd11cc25c208a']").attr('data-disqus-identifier', 'disqus_777006b154ca4af7ab8bd11cc25c208a_' + eBookConfig.course);
</script>
</div>
    </div>

</li>
<li><p class="first">Write a function to convert the image to grayscale.</p>
<blockquote>
<div>
<textarea data-component="activecode" id=ex_7_16 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true"  data-audio=''    data-gradebutton=true >


</textarea>
</div></blockquote>
</li>
<li><div id='q9' data-component="tabbedStuff" ><div data-component="tab" data-tabname="Question" >
<p>Write a function to convert an image to black and white.</p>

<textarea data-component="activecode" id=ex_7_17 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true"  data-audio=''    data-gradebutton=true >


</textarea>
</div><div data-component="tab" data-tabname="Answer" >

<textarea data-component="activecode" id=q9_answer data-lang="python"    data-timelimit=25000   data-audio=''     >
import image

def convertBlackWhite(input_image):
    grayscale_image = image.EmptyImage(input_image.getWidth(), input_image.getHeight())

    for col in range(input_image.getWidth()):
        for row in range(input_image.getHeight()):
            p = input_image.getPixel(col, row)

            red = p.getRed()
            green = p.getGreen()
            blue = p.getBlue()

            avg = (red + green + blue) / 3.0

            newpixel = image.Pixel(avg, avg, avg)
            grayscale_image.setPixel(col, row, newpixel)

    blackwhite_image = image.EmptyImage(input_image.getWidth(), input_image.getHeight())
    for col in range(input_image.getWidth()):
        for row in range(input_image.getHeight()):
            p = grayscale_image.getPixel(col, row)
            red = p.getRed()
            if red > 140:
                val = 255
            else:
                val = 0

            newpixel = image.Pixel(val, val, val)
            blackwhite_image.setPixel(col, row, newpixel)
    return blackwhite_image


win = image.ImageWin()
img = image.Image("luther.jpg")

bw_img = convertBlackWhite(img)
bw_img.draw(win)

win.exitonclick()
</textarea>
</div><div data-component="tab" data-tabname="Discussion" >
<script type="text/javascript">
    function disqus_0f0fb41d607743998a86962a11eed53d_disqus(source) { 
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'disqus_0f0fb41d607743998a86962a11eed53d_' + eBookConfig.course;
                    this.page.url = 'http://disqus_0f0fb41d607743998a86962a11eed53d_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';    
            disqus_identifier = 'disqus_0f0fb41d607743998a86962a11eed53d_' + eBookConfig.course;
            disqus_url = 'http://disqus_0f0fb41d607743998a86962a11eed53d_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'http://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class='disqus_thread_link' data-disqus-identifier="disqus_0f0fb41d607743998a86962a11eed53d" onclick="disqus_0f0fb41d607743998a86962a11eed53d_disqus(this);">Show Comments</a>
<script type='text/javascript'>
    $("a[data-disqus-identifier='disqus_0f0fb41d607743998a86962a11eed53d']").attr('data-disqus-identifier', 'disqus_0f0fb41d607743998a86962a11eed53d_' + eBookConfig.course);
</script>
</div>
    </div>

</li>
<li><p class="first">Sepia Tone images are those brownish colored images that may remind you of
times past.  The formula for creating a sepia tone is as follows:</p>
<div class="highlight-python"><div class="highlight"><pre>newR = (R × 0.393 + G × 0.769 + B × 0.189)
newG = (R × 0.349 + G × 0.686 + B × 0.168)
newB = (R × 0.272 + G × 0.534 + B × 0.131)
</pre></div>
</div>
<p>Write a function to convert an image to sepia tone. <em>Hint:</em>
Remember that rgb values must be integers between 0 and 255.</p>
<blockquote>
<div>
<textarea data-component="activecode" id=ex_7_18 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true"  data-audio=''    data-gradebutton=true >


</textarea>
</div></blockquote>
</li>
<li><div id='q11' data-component="tabbedStuff" ><div data-component="tab" data-tabname="Question" >
<p>Write a function to uniformly enlarge an image by a factor of 2 (double the size).</p>

<textarea data-component="activecode" id=ex_7_19 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true"  data-audio=''    data-gradebutton=true >


</textarea>
</div><div data-component="tab" data-tabname="Answer" >

<textarea data-component="activecode" id=answer_7_11 data-lang="python"    data-timelimit=25000   data-audio=''     >
import image

def double(oldimage):
    oldw = oldimage.getWidth()
    oldh = oldimage.getHeight()

    newim = image.EmptyImage(oldw * 2, oldh * 2)
    for row in range(oldh):
        for col in range(oldw):
            oldpixel = oldimage.getPixel(col, row)

            newim.setPixel(2*col, 2*row, oldpixel)
            newim.setPixel(2*col+1, 2*row, oldpixel)
            newim.setPixel(2*col, 2*row+1, oldpixel)
            newim.setPixel(2*col+1, 2*row+1, oldpixel)

    return newim

win = image.ImageWin()
img = image.Image("luther.jpg")

bigimg = double(img)
bigimg.draw(win)

win.exitonclick()
</textarea>
</div><div data-component="tab" data-tabname="Discussion" >
<script type="text/javascript">
    function disqus_9ca319187b4a4c2399402de0d99c0b1d_disqus(source) { 
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'disqus_9ca319187b4a4c2399402de0d99c0b1d_' + eBookConfig.course;
                    this.page.url = 'http://disqus_9ca319187b4a4c2399402de0d99c0b1d_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';    
            disqus_identifier = 'disqus_9ca319187b4a4c2399402de0d99c0b1d_' + eBookConfig.course;
            disqus_url = 'http://disqus_9ca319187b4a4c2399402de0d99c0b1d_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'http://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class='disqus_thread_link' data-disqus-identifier="disqus_9ca319187b4a4c2399402de0d99c0b1d" onclick="disqus_9ca319187b4a4c2399402de0d99c0b1d_disqus(this);">Show Comments</a>
<script type='text/javascript'>
    $("a[data-disqus-identifier='disqus_9ca319187b4a4c2399402de0d99c0b1d']").attr('data-disqus-identifier', 'disqus_9ca319187b4a4c2399402de0d99c0b1d_' + eBookConfig.course);
</script>
</div>
    </div>

</li>
<li><p class="first">After you have scaled an image too much it looks blocky.  One way of
reducing the blockiness of the image is to replace each pixel with the
average values of the pixels around it.  This has the effect of smoothing
out the changes in color.  Write a function that takes an image as a
parameter and smooths the image.  Your function should return a new image
that is the same as the old but smoothed.</p>
<blockquote>
<div>
<textarea data-component="activecode" id=ex_7_20 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true"  data-audio=''    data-gradebutton=true >


</textarea>
</div></blockquote>
</li>
<li><div id='q13' data-component="tabbedStuff" ><div data-component="tab" data-tabname="Question" >
<p>Write a general pixel mapper function that will take an image and a pixel mapping function as
parameters.  The pixel mapping function should perform a manipulation on a single pixel and return
a new pixel.</p>

<textarea data-component="activecode" id=ex_7_21 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true"  data-audio=''    data-gradebutton=true >


</textarea>
</div><div data-component="tab" data-tabname="Answer" >

<textarea data-component="activecode" id=q13_answer data-lang="python"    data-timelimit=25000   data-audio=''     >
import image

def pixelMapper(oldimage, rgbFunction):
    width = oldimage.getWidth()
    height = oldimage.getHeight()
    newim = image.EmptyImage(width, height)

    for row in range(height):
        for col in range(width):
            originalpixel = oldimage.getPixel(col, row)
            newpixel = rgbFunction(originalpixel)
            newim.setPixel(col, row, newpixel)

    return newim

def graypixel(oldpixel):
    intensitysum = oldpixel.getRed() + oldpixel.getGreen() + oldpixel.getBlue()
    aveRGB = intensitysum // 3
    newPixel = image.Pixel(aveRGB, aveRGB, aveRGB)
    return newPixel

win = image.ImageWin()
img = image.Image("luther.jpg")

newim = pixelMapper(img, graypixel)
newim.draw(win)

win.exitonclick()
</textarea>
</div><div data-component="tab" data-tabname="Discussion" >
<script type="text/javascript">
    function disqus_eb9f71a62de24efaa61f64b5a7e5d9c9_disqus(source) { 
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'disqus_eb9f71a62de24efaa61f64b5a7e5d9c9_' + eBookConfig.course;
                    this.page.url = 'http://disqus_eb9f71a62de24efaa61f64b5a7e5d9c9_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';    
            disqus_identifier = 'disqus_eb9f71a62de24efaa61f64b5a7e5d9c9_' + eBookConfig.course;
            disqus_url = 'http://disqus_eb9f71a62de24efaa61f64b5a7e5d9c9_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'http://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class='disqus_thread_link' data-disqus-identifier="disqus_eb9f71a62de24efaa61f64b5a7e5d9c9" onclick="disqus_eb9f71a62de24efaa61f64b5a7e5d9c9_disqus(this);">Show Comments</a>
<script type='text/javascript'>
    $("a[data-disqus-identifier='disqus_eb9f71a62de24efaa61f64b5a7e5d9c9']").attr('data-disqus-identifier', 'disqus_eb9f71a62de24efaa61f64b5a7e5d9c9_' + eBookConfig.course);
</script>
</div>
    </div>

</li>
<li><p class="first">When you scan in images using a scanner they may have lots of noise due to
dust particles on the image itself or the scanner itself,
or the images may even be damaged.  One way of eliminating this noise is
to replace each pixel by the median value of the pixels surrounding it.</p>
<blockquote>
<div>
<textarea data-component="activecode" id=ex_7_22 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true"  data-audio=''    data-gradebutton=true >


</textarea>
</div></blockquote>
</li>
<li><div id='q15' data-component="tabbedStuff" ><div data-component="tab" data-tabname="Question" >
<p>Research the Sobel edge detection algorithm and implement it.</p>

<textarea data-component="activecode" id=ex_7_23 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true"  data-audio=''    data-gradebutton=true >


</textarea>
</div><div data-component="tab" data-tabname="Answer" >

<textarea data-component="activecode" id=q15_answer data-lang="python"    data-timelimit=25000   data-audio=''     >
import image
import math
import sys

# Code adapted from http://www.cl.cam.ac.uk/projects/raspberrypi/tutorials/image-processing/edge_detection.html
# Licensed under the Creative Commons Attribution-ShareAlike 3.0 Unported License.

# this algorithm takes some time for larger images - this increases the amount of time
# the program is allowed to run before it times out
sys.setExecutionLimit(20000)

img = image.Image("luther.jpg")
newimg = image.EmptyImage(img.getWidth(), img.getHeight())
win = image.ImageWin()

for x in range(1, img.getWidth()-1):  # ignore the edge pixels for simplicity (1 to width-1)
    for y in range(1, img.getHeight()-1): # ignore edge pixels for simplicity (1 to height-1)

        # initialise Gx to 0 and Gy to 0 for every pixel
        Gx = 0
        Gy = 0

        # top left pixel
        p = img.getPixel(x-1, y-1)
        r = p.getRed()
        g = p.getGreen()
        b = p.getBlue()

        # intensity ranges from 0 to 765 (255 * 3)
        intensity = r + g + b

        # accumulate the value into Gx, and Gy
        Gx += -intensity
        Gy += -intensity

        # remaining left column
        p = img.getPixel(x-1, y)
        r = p.getRed()
        g = p.getGreen()
        b = p.getBlue()

        Gx += -2 * (r + g + b)

        p = img.getPixel(x-1, y+1)
        r = p.getRed()
        g = p.getGreen()
        b = p.getBlue()

        Gx += -(r + g + b)
        Gy += (r + g + b)

        # middle pixels
        p = img.getPixel(x, y-1)
        r = p.getRed()
        g = p.getGreen()
        b = p.getBlue()

        Gy += -2 * (r + g + b)

        p = img.getPixel(x, y+1)
        r = p.getRed()
        g = p.getGreen()
        b = p.getBlue()

        Gy += 2 * (r + g + b)

        # right column
        p = img.getPixel(x+1, y-1)
        r = p.getRed()
        g = p.getGreen()
        b = p.getBlue()

        Gx += (r + g + b)
        Gy += -(r + g + b)

        p = img.getPixel(x+1, y)
        r = p.getRed()
        g = p.getGreen()
        b = p.getBlue()

        Gx += 2 * (r + g + b)

        p = img.getPixel(x+1, y+1)
        r = p.getRed()
        g = p.getGreen()
        b = p.getBlue()

        Gx += (r + g + b)
        Gy += (r + g + b)

        # calculate the length of the gradient (Pythagorean theorem)
        length = math.sqrt((Gx * Gx) + (Gy * Gy))

        # normalise the length of gradient to the range 0 to 255
        length = length / 4328 * 255

        length = int(length)

        # draw the length in the edge image
        newpixel = image.Pixel(length, length, length)
        newimg.setPixel(x, y, newpixel)

newimg.draw(win)
win.exitonclick()
</textarea>
</div><div data-component="tab" data-tabname="Discussion" >
<script type="text/javascript">
    function disqus_dd2d9ca5ea744aafbf7cdc2a4ad5e974_disqus(source) { 
        if (window.DISQUS) {

            $('#disqus_thread').insertAfter(source); //put the DIV for the Disqus box after the link

            //if Disqus exists, call it's reset method with new parameters
            DISQUS.reset({
                reload: true,
                config: function () {
                    this.page.identifier = 'disqus_dd2d9ca5ea744aafbf7cdc2a4ad5e974_' + eBookConfig.course;
                    this.page.url = 'http://disqus_dd2d9ca5ea744aafbf7cdc2a4ad5e974_'+eBookConfig.course+'.interactivepython.com/#!';
                }
            });

        } else {
            //insert a wrapper in HTML after the relevant "show comments" link
            $('<div id="disqus_thread"></div>').insertAfter(source);

            // set Disqus required vars
            disqus_shortname = 'interactivepython';    
            disqus_identifier = 'disqus_dd2d9ca5ea744aafbf7cdc2a4ad5e974_' + eBookConfig.course;
            disqus_url = 'http://disqus_dd2d9ca5ea744aafbf7cdc2a4ad5e974_'+eBookConfig.course+'.interactivepython.com/#!';

            //append the Disqus embed script to HTML
            var dsq = document.createElement('script'); dsq.type = 'text/javascript'; dsq.async = true;
            dsq.src = 'http://' + disqus_shortname + '.disqus.com/embed.js';
            $('head').append(dsq);

        }
    }
</script>

<a href="#disqus_thread" class='disqus_thread_link' data-disqus-identifier="disqus_dd2d9ca5ea744aafbf7cdc2a4ad5e974" onclick="disqus_dd2d9ca5ea744aafbf7cdc2a4ad5e974_disqus(this);">Show Comments</a>
<script type='text/javascript'>
    $("a[data-disqus-identifier='disqus_dd2d9ca5ea744aafbf7cdc2a4ad5e974']").attr('data-disqus-identifier', 'disqus_dd2d9ca5ea744aafbf7cdc2a4ad5e974_' + eBookConfig.course);
</script>
</div>
    </div>

</li>
</ol>
</div>
</div>


  
      
  <li id="relations-prev" class="navLink" title='Previous Section - Glossary' data-toggle="tooltip">
    <a href="Glossary.html" >
      <i class='prevNav glyphicon glyphicon-chevron-left'></i>
    </a>
  </li>
  <a class="navLinkBg" id="navLinkBgLeft"  href="Glossary.html" ></a>
  
  <li id="relations-next" class="navLink" title='Next Section - Strings Revisited' data-toggle="tooltip" >
    <a href="../Strings/StringsRevisited.html" >
      <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right'></i>
    </a>
  </li>
  <a class="navLinkBg navLink" id="navLinkBgRight" href="../Strings/StringsRevisited.html" >Next Section - Strings Revisited</a>

<script type="text/javascript">

  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});
  
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      <span id='numuserspan'></span> readers online now | <span class='loggedinuser'></span> | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2014 Brad Miller, David Ranum, Created using Runestone Interactive.
    </p>
  </div>
</footer>



<script type="text/javascript">
  var _gaq = _gaq || [];
  _gaq.push(['_setAccount', 'UA-32029811-1']);
  _gaq.push(['_trackPageview']);

  (function() {
    var ga = document.createElement('script'); ga.type = 'text/javascript'; ga.async = true;
    ga.src = ('https:' == document.location.protocol ? 'https://ssl' : 'http://www') + '.google-analytics.com/ga.js';
    var s = document.getElementsByTagName('script')[0]; s.parentNode.insertBefore(ga, s);
  })();
</script>



  </body>
</html>