<?xml version="1.0"?>
<exercises xml:id="lists_exercises">
  <title>Exercises</title>
  <exercise>
    <statement/>
    <program interactive="" language=""/>
    <solution>
                    Your diagram should show two variables referring to two different lists.  a refers to the original list with 1,2, and 3.
                        b refers to a list with 5,2, and 3 since the zero-eth element was replaced with 5.

                    <p>Your diagram should show two variables referring to two different lists.  <c>a</c> refers to the original list with 1,2, and 3.
                        <c>b</c> refers to a list with 5,2, and 3 since the zero-eth element was replaced with 5.</p>
                </solution>
  </exercise>
  <exercise label="ex_9_2">
    <statement>
      <p>Create a list called <c>myList</c> with the following six items: 76, 92.3, <q>hello</q>, True, 4, 76.
                        Begin with the empty list shown below, and add 6 statements to add each item, one per item. The first three statements should
                        use the append method to append the item to the list, and the last three statements should use concatenation.</p>
    </statement>
    <program interactive="activecode" language="python" xml:id="ex_9_2_editor">
      <code>

myList = []

# Your code here

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(myList,[76, 92.3, 'hello', True, 4, 76],"myList should contain the specified items")
        self.assertIn(".append(", self.getEditorText(), 'append method must be used')

myTests().main()
        </code>
    </program>
    <solution>
      <program xml:id="ex_9_2_answer" language="python">
        <code>
myList = []
myList.append(76)
myList.append(92.3)
myList.append("hello")
myList = myList + [True]
myList = myList + [4]
myList = myList + [76]
print(myList)
        </code>
      </program>
    </solution>
  </exercise>
  <exercise label="ex_9_3">
    <statement>
      <p>Starting with the list of the previous exercise, write Python statements to do the following:</p>
      <p>
        <ol marker="a">
          <li>
            <p>Append <q>apple</q> and 76 to the list.</p>
          </li>
          <li>
            <p>Insert the value <q>cat</q> at position 3.</p>
          </li>
          <li>
            <p>Insert the value 99 at the start of the list.</p>
          </li>
          <li>
            <p>Find the index of <q>hello</q>.</p>
          </li>
          <li>
            <p>Count the number of 76s in the list.</p>
          </li>
          <li>
            <p>Remove the first occurrence of 76 from the list.</p>
          </li>
          <li>
            <p>Remove True from the list using <c>pop</c> and <c>index</c>.</p>
          </li>
        </ol>
      </p>
    </statement>
    <program interactive="activecode" language="python" xml:id="ex_9_3_editor">
      <code>
myList = [76, 92.3, 'hello', True, 4, 76]

# Your code here
        </code>
    </program>
    <solution>
      <program xml:id="ex_9_3_answer" language="python">
        <code>
myList = [76, 92.3, 'hello', True, 4, 76]

myList.append("apple")         # a
myList.append(76)              # a
myList.insert(3, "cat")        # b
myList.insert(0, 99)           # c

print(myList.index("hello"))   # d
print(myList.count(76))        # e
myList.remove(76)              # f
myList.pop(myList.index(True)) # g

print (myList)
        </code>
      </program>
    </solution>
  </exercise>
  <exercise label="ex_9_4">
    <statement>
      <p>Write a function called <c>average</c> that takes a list of numbers as a parameter and returns the average of the numbers.</p>
    </statement>
    <program interactive="activecode" language="python" xml:id="ex_9_4_editor">
      <code>
def average(numlist):
    # Complete the function definition

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(int(average([1, 3, 5, 7])),4,"average([1, 3, 5, 7]) should be 4")
        self.assertEqual(int(average([9, 5, 4])),6,"average([9, 5, 4]) should be 6")

myTests().main()
        </code>
    </program>
    <solution>
      <program xml:id="ex_9_4_solution" language="python">
        <code>
def average(numlist):

    total = 0
    for num in numlist:
        total = total + num

    return total / len(numlist)
        </code>
      </program>
    </solution>
  </exercise>
  <exercise label="ex_9_5">
    <statement>
      <p>Write a Python function named <c>max</c> that takes a parameter containing a nonempty list of integers and returns the maximum value.  (Note: there is a builtin function named <c>max</c> but pretend you cannot use it.)</p>
    </statement>
    <program interactive="activecode" language="python" xml:id="ex_9_5_editor">
      <code>
def max(lst):
    # Complete the function

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(max([3, 31, 5, 7]),31,"max([3, 31, 5, 7]) should be 31")
        self.assertEqual(max([3, 13, 51, 7]),51,"max([3, 13, 51, 7]) should be 51")

myTests().main()
        </code>
    </program>
    <solution>
      <program xml:id="lst_q5_answer" language="python">
        <code>
def max(lst):
    max = lst[0]
    for e in lst:
        if e &gt; max:
            max = e
    return max
        </code>
      </program>
    </solution>
  </exercise>
  <exercise label="ex_7_11">
    <statement>
      <p>Write a function <c>sum_of_squares(xs)</c> that computes the sum
                of the squares of the numbers in the list <c>xs</c>.  For example,
                <c>sum_of_squares([2, 3, 4])</c> should return 4+9+16 which is 29:</p>
    </statement>
    <program xml:id="ex_7_11_editor" interactive="activecode" language="python">
      <code>
def sum_of_squares(xs):
    # your code here

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(sum_of_squares([2,3,4]),29,"Tested sum_of_squares on input [2,3,4]")
        self.assertEqual(sum_of_squares([0,1,-1]),2,"Tested sum_of_squares on input [0,1,-1]")
        self.assertEqual(sum_of_squares([5,12,14]),365,"Tested sum_of_squares on input [5,12,14]")

myTests().main()
        </code>
    </program>
  </exercise>
  <exercise label="ex_9_6">
    <statement>
      <p>Write a function to count how many odd numbers are in a list.</p>
    </statement>
    <program interactive="activecode" language="python" xml:id="ex_9_6_editor">
      <code>
def countOdd(lst):
    # your code here

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(countOdd([1,3,5,7,9]),5,"Tested countOdd on input [1,3,5,7,9]")
        self.assertEqual(countOdd([1,2,3,4,5]),3,"Tested countOdd on input [-1,-2,-3,-4,-5]")
        self.assertEqual(countOdd([2,4,6,8,10]),0,"Tested countOdd on input [2,4,6,8,10]")
        self.assertEqual(countOdd([0,-1,12,-33]),2,"Tested countOdd on input [0,-1,12,-33]")

myTests().main()
        </code>
    </program>
    <solution>
      <program xml:id="lst_q7_answer" language="python">
        <code>
import random

def countOdd(lst):
    odd = 0
    for e in lst:
        if e % 2 != 0:
            odd = odd + 1
    return odd

# make a random list to test the function
lst = []
for i in range(100):
    lst.append(random.randint(0, 1000))

print(countOdd(lst))
        </code>
      </program>
    </solution>
  </exercise>
  <exercise label="ex_9_7">
    <statement>
      <p>Sum up all the even numbers in a list.</p>
    </statement>
    <program xml:id="ex_9_7_editor" interactive="activecode" language="python">
      <code>
def sumEven(lst):
    # your code here

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(sumEven([1,3,5,7,9]),0,"Tested sumEven on input [1,3,5,7,9]")
        self.assertEqual(sumEven([-1,-2,-3,-4,-5]),-6,"Tested sumEven on input [-1,-2,-3,-4,-5]")
        self.assertEqual(sumEven([2,4,6,7,9]),12,"Tested sumEven on input [2,4,6,7,9]")
        self.assertEqual(sumEven([0,1,12,33]),12,"Tested sumEven on input [0,1,12,33]")

myTests().main()
        </code>
    </program>
  </exercise>
  <exercise label="ex_9_8">
    <statement>
      <p>Sum up all the negative numbers in a list.</p>
    </statement>
    <program interactive="activecode" language="python" xml:id="ex_9_8_editor">
      <code>
def sumNegatives(lst):
    # your code here

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(sumNegatives([-1,-2,-3,-4,-5]),-15,"Tested sumNegatives on input [-1,-2,-3,-4,-5]")
        self.assertEqual(sumNegatives([1,-3,5,-7,9]),-10,"Tested sumNegatives on input [1,-3,5,-7,9]")
        self.assertEqual(sumNegatives([-2,-4,6,-7,9]),-13,"Tested sumNegatives on input [-2,-4,6,-7,9]")
        self.assertEqual(sumNegatives([0,1,2,3,4]),0,"Tested sumNegatives on input [0,1,2,3,4]")

myTests().main()
        </code>
    </program>
    <solution>
      <program xml:id="lst_q9_answer" language="python">
        <code>
import random

def sumNegative(lst):
    sum = 0
    for e in lst:
        if e &lt; 0:
            sum = sum + e
    return sum

lst = []
for i in range(100):
    lst.append(random.randrange(-1000, 1000))

print(sumNegative(lst))
        </code>
      </program>
    </solution>
  </exercise>
  <exercise label="ex_9_9">
    <statement>
      <p>Count how many words in a list have length 5.</p>
    </statement>
    <program xml:id="ex_9_9_editor" interactive="activecode" language="python">
      <code>
def countWords(lst):
    # your code here
        </code>
    </program>
  </exercise>
  <exercise>
    <statement/>
    <program interactive="" language=""/>
    <solution>
      <program xml:id="lst_q11_answer" language="python">
        <code>
import random

def sum(lst):
    sum = 0
    index = 0
    while index &lt; len(lst) and lst[index] % 2 != 0:
        sum = sum + lst[index]
        index = index + 1
    return sum

lst = []
for i in range(100):
    lst.append(random.randint(0,1000))

print(sum(lst))
        </code>
      </program>
    </solution>
  </exercise>
  <exercise label="ex_9_11">
    <statement>
      <p>Count how many words occur in a list up to and including the first occurrence of the word <q>sam</q>.</p>
    </statement>
    <program xml:id="ex_9_11_editor" interactive="activecode" language="python">
      <code>
def count(lst):
    # your code here
        </code>
    </program>
  </exercise>
  <exercise label="ex_9_12">
    <statement>
      <p>Although Python provides us with many list methods, it is good practice and very instructive to think about how they are implemented.  Implement a Python function that works like the following:</p>
      <p>
        <ol marker="a">
          <li>
            <p>count</p>
          </li>
          <li>
            <p>in</p>
          </li>
          <li>
            <p>reverse</p>
          </li>
          <li>
            <p>index</p>
          </li>
          <li>
            <p>insert</p>
          </li>
        </ol>
      </p>
    </statement>
    <program interactive="activecode" language="python" xml:id="ex_9_12_editor">
      <code>

        </code>
    </program>
    <solution>
      <program xml:id="lst_q13_answer" language="python">
        <code>
def count(obj, lst):
    count = 0
    for e in lst:
        if e == obj:
            count = count + 1
    return count

def is_in(obj, lst):  # cannot be called in() because in is a reserved keyword
    for e in lst:
        if e == obj:
            return True
    return False

def reverse(lst):
    reversed = []
    for i in range(len(lst)-1, -1, -1): # step through the original list backwards
        reversed.append(lst[i])
    return reversed

def index(obj, lst):
    for i in range(len(lst)):
        if lst[i] == obj:
            return i
    return -1

def insert(obj, index, lst):
    newlst = []
    for i in range(len(lst)):
        if i == index:
            newlst.append(obj)
        newlst.append(lst[i])
    return newlst

lst = [0, 1, 1, 2, 2, 3, 4, 5, 6, 7, 8, 9]
print(count(1, lst))
print(is_in(4, lst))
print(reverse(lst))
print(index(2, lst))
print(insert('cat', 4, lst))
        </code>
      </program>
    </solution>
  </exercise>
  <exercise label="ex_9_13">
    <statement>
      <p>Write a function <c>replace(s, old, new)</c> that replaces all occurences of
                <c>old</c> with <c>new</c> in a string <c>s</c>:</p>
      <pre>test(replace('Mississippi', 'i', 'I'), 'MIssIssIppI')

s = 'I love spom!  Spom is my favorite food.  Spom, spom, spom, yum!'
test(replace(s, 'om', 'am'),
       'I love spam!  Spam is my favorite food.  Spam, spam, spam, yum!')

test(replace(s, 'o', 'a'),
       'I lave spam!  Spam is my favarite faad.  Spam, spam, spam, yum!')</pre>
      <p><em>Hint</em>: use the <c>split</c> and <c>join</c> methods.</p>
    </statement>
    <program xml:id="ex_9_13_editor" interactive="activecode" language="python">
      <code>
def replace(s, old, new):
    # your code here

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(replace('Mississippi','i','I'),'MIssIssIppI',"Tested replace on input 'Mississippi','i','I'")
        self.assertEqual(replace('Bookkeeper','e','A'),'BookkAApAr',"Tested failed on input 'Bookkeeper','e','A'")
        self.assertEqual(replace('Deeded','e','q'),'Dqqdqd',"Tested failed on input 'Deeded','e','q'")

myTests().main()
        </code>
    </program>
  </exercise>
  <exercise label="ex_9_14">
    <statement>
      <p>Here are the rules for an L-system that creates something that resembles
                        a common garden herb.  Implement the following rules and try it.  Use an
                        angle of 25.7 degrees.</p>
      <pre>H
H --&gt; HFX[+H][-H]
X --&gt; X[-FFF][+FFF]FX</pre>
    </statement>
    <program interactive="activecode" language="python" xml:id="ex_9_14_editor">
      <code>

        </code>
    </program>
    <solution>
      <program xml:id="lst_q15_answer" language="python">
        <code>
import turtle

def createLSystem(numIters, axiom):
    startString = axiom
    endString = ""
    for i in range(numIters):
        endString = processString(startString)
        startString = endString

    return endString

def processString(oldStr):
    newstr = ""
    for ch in oldStr:
        newstr = newstr + applyRules(ch)

    return newstr

def applyRules(ch):
    newstr = ""
    if ch == 'H':
        newstr = 'HFX[+H][-H]'   # Rule 1
    elif ch == 'X':
        newstr = 'X[-FFF][+FFF]FX'
    else:
        newstr = ch     # no rules apply so keep the character

    return newstr

def drawLsystem(aTurtle, instructions, angle, distance):
    savedInfoList = []
    for cmd in instructions:
        if cmd == 'F':
            aTurtle.forward(distance)
        elif cmd == 'B':
            aTurtle.backward(distance)
        elif cmd == '+':
            aTurtle.right(angle)
        elif cmd == '-':
            aTurtle.left(angle)
        elif cmd == '[':
            savedInfoList.append([aTurtle.heading(), aTurtle.xcor(), aTurtle.ycor()])
            #print(savedInfoList)
        elif cmd == ']':
            newInfo = savedInfoList.pop()
            aTurtle.setheading(newInfo[0])
            aTurtle.setposition(newInfo[1], newInfo[2])


def main():
    inst = createLSystem(4, "H")   # create the string
    print(inst)
    t = turtle.Turtle()            # create the turtle
    wn = turtle.Screen()
    t.up()
    t.back(200)
    t.down()
    t.speed(9)
    drawLsystem(t, inst, 27.5, 5)  # draw the picture

    wn.exitonclick()

main()
        </code>
      </program>
    </solution>
  </exercise>
  <exercise label="ex_9_16">
    <statement>
      <p>Here is another L-System.  Use an Angle of 25.</p>
      <pre>F
F --&gt; F[-F]F[+F]F</pre>
    </statement>
    <program xml:id="ex_9_16_editor" interactive="activecode" language="python">
      <code>

        </code>
    </program>
  </exercise>
  <exercise label="ex_9_17">
    <statement>
      <p>Create a list named <c>randlist</c> containing 100 random integers between 0 and 1000 (use iteration, append, and the random module).</p>
    </statement>
    <program interactive="activecode" language="python" xml:id="ex_9_17_editor">
      <code>

====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(len(randlist),100,"randlist should contain 100 numbers")

myTests().main()
        </code>
    </program>
    <solution>
      <program xml:id="ac_ex_9_17" language="python">
        <code>
import random

randlist = []
for i in range(100):
    randlist.append(random.randint(0, 1000))
        </code>
      </program>
    </solution>
  </exercise>
</exercises>
